# Cadenas de texto (Strings)
 
Temas:
1. C贸mo obtener una cadena de texto
2. Interpolaci贸n de cadenas
3. Concatenaci贸n de cadenas

## C贸mo obtener una cadena de texto
 
隆Encierra tus caracteres entre " " o """ """!

```julia
s1 = "I am a string."
```

Output:
```
"I am a string."
```

```julia
s2 = """I am also a string. """
```

Output:
```
"I am also a string. "
```

Existen un par de diferencias funcionales entre las cadenas encerradas en comillas simples y triples. <br>
Una diferencia es que, en el segundo caso, puedes usar comillas dentro de tu cadena.

```julia
"Here, we get an "error" because it's ambiguous where this string ends "
```

Output:
```
@ ./boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1116
```

```julia
"""Look, Mom, no "errors"!!! """
```

Output:
```
"Look, Mom, no \"errors\"!!! "
```

隆Nota que ' ' define un car谩cter, pero NO una cadena de texto!

```julia
typeof('a')
```

Output:
```
Char
```

```julia
'We will get an error here'
```

Output:
```
@ ./boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1116
```

## Interpolaci贸n de cadenas
 
Podemos usar el signo $ para insertar variables existentes en una cadena y para evaluar expresiones dentro de una cadena. <br>
Abajo hay un ejemplo que contiene informaci贸n personal altamente sensible.

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

Output:
```
Hello, my name is Jane.
I have 10 fingers and 10 toes.
```

```julia
 println("That is $(num_fingers + num_toes) digits in all!!")
```

Output:
```
That is 20 digits in all!!
```

## Concatenaci贸n de cadenas
 
隆Abajo hay tres formas de concatenar cadenas! <br><br>
La primera forma es usar la funci贸n `string()`. <br>
`string()` convierte entradas que no son cadenas a cadenas.

```julia
s3 = "How many cats ";
s4 = "is too many cats?";
 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"How many cats is too many cats?"
```

```julia
string("I don't know, but ", , " is too few.")
```

Output:
```
"I don't know, but 10 is too few."
```

隆Tambi茅n podemos usar `*` para concatenar!

```julia
s3*s4
```

Output:
```
"How many cats is too many cats?"
```

### Ejercicios
 
#### 2.1 
Crea una cadena que diga "hi" 1000 veces, primero con `repeat` y luego con el operador de potenciaci贸n, que puede llamar a `*` internamente. As铆gnala a la variable `hi` abajo.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia

```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

#### 2.2 
Declara dos variables
 
```julia
a = 3
b = 4
```
y 煤salas para crear dos cadenas:
```julia
"3 + 4"
"7" 
```
y guarda los resultados en `c` y `d` respectivamente

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia

```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"
```