# Álgebra lineal básica en Julia
Autor: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(con ediciones de Jane Herriman)

Primero definamos una matriz aleatoria

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 3  1  3
 1  1  2
 3  2  2
```

Definir un vector de unos

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Observa que $A$ tiene tipo Array{Int64,2} pero $x$ tiene tipo Array{Float64,1}. Julia define los alias Vector{Type}=Array{Type,1} y Matrix{Type}=Array{Type,2}. 

Muchas de las operaciones básicas son las mismas que en otros lenguajes
#### Multiplicación

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 7.0
 4.0
 7.0
```

#### Transposición
Como en otros lenguajes `A'` es la transpuesta conjugada, o adjunta

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 3  1  3
 1  1  2
 3  2  2
```

y podemos obtener la transpuesta con

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 3  1  3
 1  1  2
 3  2  2
```

#### Multiplicación transpuesta
Julia nos permite escribir esto sin *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 19  10  17
 10   6   9
 17   9  17
```

#### Resolución de sistemas lineales 
El problema $Ax=b$ para $A$ ***cuadrada*** se resuelve con la función \.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

`A\b` nos da la *solución de mínimos cuadrados* si tenemos un sistema lineal sobredeterminado (una matriz "alta")

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.01259    0.562346
 0.571931   0.208287
 0.0761351  0.507003
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
  2.3722998636787898
 12.85556853172185
```

y la *solución de mínimos cuadrados de norma mínima* si tenemos un problema de mínimos cuadrados deficiente en rango

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.682093  0.682093
 0.745287  0.745287
 0.698837  0.698837
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 4.190527048390038
 4.190527048390039
```

Julia también nos da la solución de norma mínima cuando tenemos una solución subdeterminada (una matriz "corta")

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.692218  0.193236   0.153774
 0.127852  0.0359394  0.063034
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -4.139674407706176
 -0.9797338415813472
 23.563074056363842
```

# La librería LinearAlgebra

Aunque mucho del álgebra lineal está disponible en Julia por defecto (como se mostró arriba), hay una librería estándar llamada `LinearAlgebra` que incorpora muchos más nombres y funciones relevantes. En particular, proporciona factorizaciones y algunos tipos de matrices estructuradas. Como con todos los paquetes, puedes traer estas características adicionales a tu sesión con `using LinearAlgebra`.

### Ejercicios

#### 10.1 
Toma el producto interior (o producto "punto") de un vector `v` consigo mismo y asígnalo a la variable `dot_v`.

```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Toma el producto exterior de un vector v consigo mismo y asígnalo a la variable `outer_v`

```julia

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Usa [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) para calcular el producto cruz de un vector v consigo mismo y asígnalo a la variable `cross_v`

```julia

```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```