# Estructuras de datos
 
Cuando empezamos a trabajar con muchos datos a la vez, es conveniente almacenarlos en estructuras como arreglos o diccionarios (en lugar de depender solo de variables).<br>
 
Tipos de estructuras de datos que veremos:
1. Tuplas
2. Diccionarios
3. Arreglos
 
<br>
Como resumen, las tuplas y los arreglos son secuencias ordenadas de elementos (por lo que podemos indexarlas). Los diccionarios y los arreglos son mutables.
¡Explicaremos esto más abajo!

## Tuplas
 
Podemos crear una tupla encerrando una colección ordenada de elementos entre `( )`.
 
Sintaxis: <br>
```julia
(elemento1, elemento2, ...)
```
 

```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Output:
```
("penguins", "cats", "sugargliders")
```

Podemos indexar esta tupla,

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

pero como las tuplas son inmutables, no podemos actualizarlas

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
@ ./boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1094
```

## Nuevo en 1.6: NamedTuples
 
Como puedes imaginar, los `NamedTuple` son como las `Tuple` pero cada elemento tiene además un nombre. Tienen una sintaxis especial usando `=` dentro de la tupla:
 
```julia
(nombre1 = elemento1, nombre2 = elemento2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Al igual que las `Tuples` normales, los `NamedTuples` son ordenados, por lo que podemos recuperar sus elementos mediante índices:

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

También añaden la habilidad especial de acceder a los valores por su nombre:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

## Diccionarios
 
Si tenemos conjuntos de datos relacionados entre sí, podemos elegir almacenarlos en un diccionario. Podemos crear un diccionario usando la función `Dict()`, que podemos inicializar como un diccionario vacío o con pares clave-valor.
 
Sintaxis:
```julia
Dict(clave1 => valor1, clave2 => valor2, ...)
```
 
Un buen ejemplo es una lista de contactos, donde asociamos nombres con números de teléfono.

```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

En este ejemplo, cada nombre y número es un par "clave" y "valor". Podemos obtener el número de Jenny (un valor) usando la clave asociada

```julia
myphonebook["Jenny"]
```

Output:
```
"867-5309"
```

Podemos agregar otra entrada a este diccionario así

```julia
myphonebook["Kramer"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Veamos cómo luce ahora nuestra agenda telefónica...

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555-FILK"
  "Ghostbusters" => "555-2368"
```

Podemos eliminar a Kramer de nuestra lista de contactos —y al mismo tiempo obtener su número— usando `pop!`

```julia
pop!(myphonebook, "Kramer")
```

Output:
```
"555-FILK"
```

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

A diferencia de las tuplas y los arreglos, los diccionarios no son ordenados. Por lo tanto, no podemos indexarlos.

```julia
myphonebook[1]
```

Output:
```
@ ./boot.jl:360 [inlined]
 [4] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1094
```

En el ejemplo anterior, `julia` piensa que intentas acceder a un valor asociado a la clave `1`.

## Arreglos
 
A diferencia de las tuplas, los arreglos son mutables. A diferencia de los diccionarios, los arreglos contienen colecciones ordenadas. <br>
Podemos crear un arreglo encerrando la colección entre `[ ]`.
 
Sintaxis: <br>
```julia
[elemento1, elemento2, ...]
```
 
Por ejemplo, podríamos crear un arreglo para llevar registro de mis amigos

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]
```

Output:
```
5-element Vector{String}:
 "Ted"
 "Robyn"
 "Barney"
 "Lily"
 "Marshall"
```

El `1` en `Array{String,1}` significa que es un vector unidimensional. Un `Array{String,2}` sería una matriz 2D, etc. `String` es el tipo de cada elemento.

o para almacenar una secuencia de números

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Ted"
  "Robyn"
```

Una vez que tenemos un arreglo, podemos obtener elementos individuales indexando el arreglo. Por ejemplo, si queremos el tercer amigo listado en `myfriends`, escribimos

```julia
myfriends[3]
```

Output:
```
"Barney"
```

Podemos usar la indexación para editar un elemento existente de un arreglo

```julia
myfriends[3] = "Baby Bop"
```

Output:
```
"Baby Bop"
```

Sí, Julia usa indexación desde 1, no desde 0 como Python. Se han librado guerras por cosas menores. Tengo un amigo con la sabiduría de Salomón que propone zanjar esto de una vez por todas con ½ 😃

También podemos editar el arreglo usando las funciones `push!` y `pop!`. `push!` agrega un elemento al final del arreglo y `pop!` elimina el último elemento.
 
Podemos agregar otro número a nuestra secuencia de Fibonacci

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
 21
```

y luego eliminarlo

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

Hasta ahora solo he dado ejemplos de arreglos 1D de escalares, pero los arreglos pueden tener cualquier número de dimensiones y también pueden almacenar otros arreglos. 
<br><br>
Por ejemplo, los siguientes son arreglos de arreglos:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["koobideh", "chocolate", "eggs"]
 ["penguins", "cats", "sugargliders"]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Abajo hay ejemplos de arreglos 2D y 3D poblados con valores aleatorios.

```julia
rand(4, 3)
```

Output:
```
4×3 Matrix{Float64}:
 0.392655  0.340897  0.736004
 0.887521  0.625563  0.225457
 0.123646  0.215211  0.435754
 0.563704  0.564539  0.0301484
```

```julia
rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.855458  0.180741   0.480692
 0.888806  0.852159   0.62636
 0.708986  0.0526796  0.779808
 0.499323  0.18727    0.0843665

[:, :, 2] =
 0.641847  0.292116  0.0178846
 0.359826  0.9608    0.245789
 0.746938  0.389501  0.944762
 0.386954  0.391897  0.0875665
```

¡Ten cuidado cuando quieras copiar arreglos!

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers = fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
 404
   1
   2
   3
   5
   8
  13
```

¡Editar `somenumbers` hizo que `fibonacci` también se actualizara!
 
En el ejemplo anterior, en realidad no hicimos una copia de `fibonacci`. Solo creamos una nueva forma de acceder a las entradas del arreglo referenciado por `fibonacci`.
 
Si queremos hacer una copia del arreglo referenciado por `fibonacci`, podemos usar la función `copy`.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somemorenumbers = copy(fibonacci)
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
somemorenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

En este último ejemplo, fibonacci no se actualizó. Por lo tanto, vemos que los arreglos referenciados por `somemorenumbers` y `fibonacci` son distintos.

### Ejercicios
 
#### 3.1 
Crea un arreglo, `a_ray`, con el siguiente código:
 
```julia
a_ray = [1, 2, 3]
```
 
Agrega el número `4` al final de este arreglo y luego elimínalo.

```julia

```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Intenta agregar "Emergency" como clave a `myphonebook` con el valor `string(911)` con el siguiente código
```julia
myphonebook["Emergency"] = 911
```
 
¿Por qué no funciona esto?

```julia

```

#### 3.3 
Crea un nuevo diccionario llamado `flexible_phonebook` que tenga el número de Jenny almacenado como entero y el de los Cazafantasmas como cadena, con el siguiente código
 
```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Agrega la clave "Emergency" con el valor `911` (un entero) a `flexible_phonebook`.

```julia

```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
¿Por qué podemos agregar un entero como valor a `flexible_phonebook` pero no a `myphonebook`? ¿Cómo podríamos haber inicializado `myphonebook` para que aceptara enteros como valores? (pista: consulta la [documentación de diccionarios de Julia](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

```julia

```