# Bucles (Loops)
 
Temas:
1. Bucles `while`
2. Bucles `for`
<br>
 
## Bucles while
 
La sintaxis para un `while` es:
 
```julia
while *condición*
    *cuerpo del bucle*
end
```
 
Por ejemplo, podríamos usar `while` para contar o para iterar sobre un arreglo.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## Bucles for
 
La sintaxis para un bucle `for` es:
 
```julia
for *var* in *colección*
    *cuerpo del bucle*
end
```
 
Podríamos usar un bucle for para generar los mismos resultados que cualquiera de los ejemplos anteriores:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Ahora usemos bucles `for` para crear algunas tablas de suma, donde el valor de cada entrada es la suma de sus índices de fila y columna. <br>
 
Nota que iteramos sobre este arreglo con bucles en orden columna-primero para obtener el mejor rendimiento. Más información sobre el acceso rápido a arreglos multidimensionales en bucles anidados se puede encontrar en https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1
 
Primero, inicializamos un arreglo con ceros.

```julia
m, n = 5, 5
A = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

Aquí tienes una forma más concisa (azúcar sintáctica) para el mismo bucle `for` anidado

```julia
B = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

La forma más "juliana" de crear esta tabla de suma habría sido con una *comprensión de arreglos* (array comprehension).

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

### Ejercicios
 
#### 4.1 
Itera sobre los enteros del 1 al 100 e imprime sus cuadrados.

```julia

```

#### 4.2 
Agrega al código anterior un poco para crear un diccionario, `squares`, que almacene enteros y sus cuadrados como pares clave-valor de modo que
 
```julia
squares[10] == 100
```

```julia

```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Usa una comprensión de arreglos para crear un arreglo `squares_arr` que almacene los cuadrados de todos los enteros del 1 al 100.

```julia

```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```