# Funciones
 
Temas:
1. Cómo declarar una función
2. Duck-typing en Julia
3. Funciones mutantes vs. no mutantes
4. Algunas funciones de orden superior

## Cómo declarar una función
Julia nos da varias formas de escribir una función. La primera requiere las palabras clave `function` y `end`

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

Output:
```
sayhi (generic function with 1 method)
```

```julia
function f(x)
    x^2
end
```

Output:
```
f (generic function with 1 method)
```

Podemos llamar a cualquiera de estas funciones así:

```julia
sayhi("C-3PO")
```

Output:
```
Hi C-3PO, it's great to see you!
```

```julia
f(42)
```

Output:
```
1764
```

Alternativamente, podríamos haber declarado cualquiera de estas funciones en una sola línea

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

Output:
```
sayhi2 (generic function with 1 method)
```

```julia
f2(x) = x^2
```

Output:
```
f2 (generic function with 1 method)
```

```julia
sayhi2("R2D2")
```

Output:
```
Hi R2D2, it's great to see you!
```

```julia
f2(42)
```

Output:
```
1764
```

Finalmente, podríamos haber declarado estas como funciones "anónimas"

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

Output:
```
#1 (generic function with 1 method)
```

```julia
f3 = x -> x^2
```

Output:
```
#3 (generic function with 1 method)
```

```julia
sayhi3("Chewbacca")
```

Output:
```
Hi Chewbacca, it's great to see you!
```

```julia
f3(42)
```

Output:
```
1764
```

## Duck-typing en Julia
*"Si grazna como un pato, es un pato."* <br><br>
Las funciones en Julia funcionarán con cualquier entrada que tenga sentido. <br><br>
Por ejemplo, `sayhi` funciona con el nombre de este personaje menor de TV, escrito como un entero...

```julia
sayhi(55595472)
```

Output:
```
Hi 55595472, it's great to see you!
```

Y `f` funcionará con una matriz.

```julia
A = rand(3, 3)
A
```

Output:
```
3×3 Matrix{Float64}:
 0.983045  0.336857  0.21284
 0.102728  0.711431  0.817376
 0.504421  0.681324  0.833934
```

```julia
f(A)
```

Output:
```
3×3 Matrix{Float64}:
 1.10834   0.715809  0.662064
 0.586371  1.09764   1.28501
 0.986512  1.22281   1.3597
```

`f` también funcionará con una cadena como "hi" porque `*` está definido para entradas de tipo string como concatenación de cadenas.

```julia
f("hi")
```

Output:
```
"hihi"
```

Por otro lado, `f` no funcionará con un vector. A diferencia de `A^2`, que está bien definido, el significado de `v^2` para un vector `v` no es una operación algebraica bien definida.

```julia
v = rand(3)
```

Output:
```
3-element Vector{Float64}:
 0.1488599996403983
 0.5795662436976401
 0.12786532477145474
```

```julia
# This won't work
f(v)
```

Output:
```
@ ./boot.jl:360 [inlined]
 [6] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1094
```

## Funciones mutantes vs. no mutantes
 
Por convención, las funciones que terminan con `!` alteran su contenido y las que no tienen `!` no lo hacen.
 
Por ejemplo, veamos la diferencia entre `sort` y `sort!`.

```julia
v = [3, 5, 2]
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

```julia
sort(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

`sort(v)` devuelve un arreglo ordenado que contiene los mismos elementos que `v`, pero `v` permanece sin cambios. <br><br>
 
En cambio, cuando ejecutamos `sort!(v)`, el contenido de `v` se ordena dentro del mismo arreglo.

```julia
sort!(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

## Algunas funciones de orden superior
 
### map
 
`map` es una función de "orden superior" en Julia que *toma una función* como uno de sus argumentos de entrada. 
`map` aplica esa función a cada elemento de la estructura de datos que le pases. Por ejemplo, ejecutar
 
```julia
map(f, [1, 2, 3])
```
te dará un arreglo de salida donde la función `f` se ha aplicado a todos los elementos de `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Aquí hemos elevado al cuadrado todos los elementos del vector `[1, 2, 3]`, en lugar de elevar al cuadrado el vector `[1, 2, 3]`.
 
Para esto, podríamos haber pasado a `map` una función anónima en vez de una función con nombre, como

```julia
x -> x^3
```

Output:
```
#5 (generic function with 1 method)
```

¡Eso es todo por ahora sobre funciones!
 
En la siguiente sección, aprenderemos sobre paquetes en Julia.

```julia
map(x -> x^3, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
  1
  8
 27
```

¿Preguntas o dudas sobre funciones en Julia? ¡No dudes en preguntar!

## Resumen
- Las funciones en Julia se definen con la palabra clave `function` o con una sintaxis abreviada.
- Se pueden pasar funciones como argumentos a otras funciones.
- Las funciones anónimas son útiles para operaciones rápidas y sencillas.

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

### Ejercicios propuestos
1. Escribe una función que reciba un número y devuelva su valor absoluto.
2. Escribe una función que reciba un vector y devuelva el promedio de sus elementos.
3. Escribe una función que reciba dos números y devuelva el mayor de ellos.
4. Escribe una función anónima que multiplique por 10 cada elemento de un vector.
5. Usa `map` para aplicar la función del punto anterior a `[1, 2, 3, 4, 5]`.

```julia
f.([1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

¿Listo para seguir? ¡Vamos a aprender sobre paquetes en Julia!

## Soluciones a los ejercicios
1. 
```julia
function valor_absoluto(x)
    return abs(x)
end
```
2. 
```julia
function promedio(v)
    return sum(v) / length(v)
end
```
3. 
```julia
function mayor(a, b)
    return a > b ? a : b
end
```
4. 
```julia
multiplicar_por_10 = x -> x * 10
```
5. 
```julia
map(multiplicar_por_10, [1, 2, 3, 4, 5])
```

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Matrix{Int64}:
  30   36   42
  66   81   96
 102  126  150
```

¡Felicidades! Ahora sabes cómo definir y usar funciones en Julia, incluyendo funciones anónimas y cómo pasarlas como argumentos.

```julia
B = f.(A)
```

Output:
```
3×3 Matrix{Int64}:
  1   4   9
 16  25  36
 49  64  81
```

¿Quieres practicar más? Intenta crear funciones para resolver problemas cotidianos, como calcular el área de un círculo o convertir temperaturas de Celsius a Fahrenheit.

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

¡Nos vemos en la siguiente lección!

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

---

# Recursos adicionales
- [Documentación oficial de Julia: Funciones](https://docs.julialang.org/en/v1/manual/functions/)
- [Más ejemplos de funciones en Julia](https://juliabyexample.helpmanual.io/)
- [Funciones anónimas en Julia](https://docs.julialang.org/en/v1/manual/functions/#man-anonymous-functions)

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Use `map` or `broadcast` to increment every element of matrix `A` by `1` and assign it to a variable `A1`.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 
Use the broadcast dot syntax to increment every element of matrix `A1` by `1` and store it in variable `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```